---
title_meta  : Chapter 1
title       : Insert the chapter title here
description : What is this chapter about? Add here your description

---
## Create Courses for DataCamp

testtest

*** =instructions
- instr

*** =hint
asdf

*** =pre_exercise_code
```{r,eval=FALSE}
x = 4
```

*** =sample_code
```{r,eval=FALSE}
y = 5
```

*** =solution
```{r,eval=FALSE}
c(1,3,3)
```

*** =sct
```{r,eval=FALSE}
# The Submission Correctness Test code
# It takes as input the students workspace, his code and what was printed to the console during the exercise.
# Based on these inputs the SCT returns:
# (i) whether the students code was correct
# (ii) a message to the student to either help him find what was wrong or congratulate him on finding the solution.
success_msg("Well done! Continue to see some more examples!")
```